```python
# Initialize Otter
import otter
grader = otter.Notebook("longshort_strategy.ipynb")
```

```python
import datetime as dt
import pandas as pd
import numpy as np
import warnings
import statsmodels.api as sm
import matplotlib.pyplot as plt
from scipy.stats import zscore
from sklearn.metrics import mean_squared_error, r2_score

warnings.filterwarnings('ignore')

def isclose(value, original, tolerance= 0.05):
    return value <= original * (1+tolerance) and value >= original * (1-tolerance)
```

In this homework, we will:
- Introduce pandas through data cleaning and processing.
- Implement a basic trading strategy.
- Model the costs of a trading strategy.
- Calculate expected returns.
- Use expected returns to create a trading strategy.
- Try to add predictive features/signals.

# Pandas intro

```python
!pip install lxml
```

Output:
```
Collecting lxml
  Using cached lxml-5.3.2-cp311-cp311-manylinux_2_28_x86_64.whl.metadata (3.6 kB)
Using cached lxml-5.3.2-cp311-cp311-manylinux_2_28_x86_64.whl (5.0 MB)
Installing collected packages: lxml
Successfully installed lxml-5.3.2
```

```python
# list of sp500 companies, from wikipedia
url = 'https://en.m.wikipedia.org/w/index.php?title=List_of_S%26P_500_companies&oldid=1219339331'

# pandas lets us read html tables from wikipedia page
data = pd.read_html(url)
```

You can add cells in the notebook to perform quick operations or tests. Note that the last value in a cell will automatically be displayed below the cell once it finishes running. Let's examine the type of variable that `data` returns.

```python
type(data)
```

Output:
```
list
```

We can access values in a list using indexing and slicing operations:

- `list[i]` will return the \(i\)-th value in the list, starting with 0.
- `list[i:]` will return the \(i\)-th value and onward.
- `list[:i]` will return all values up to the \(i\)-th value.
- `list[i:j]` will return values from the \(i\)-th value onward, up to but not including the \(j\)-th value.

```python
# looks like data returns a list of tables, what type is each table? 
type(data[0])
```

Output:
```
pandas.core.frame.DataFrame
```

```python
# the dataframe is the core object of pandas, we'll abbreviate as df
# a dataframe is essentially a code version of a table
df = data[0]
df
```

Output:
```
Symbol            Security             GICS Sector  \
0      MMM                  3M             Industrials   
1      AOS         A. O. Smith             Industrials   
2      ABT              Abbott             Health Care   
3     ABBV              AbbVie             Health Care   
4      ACN           Accenture  Information Technology   
..     ...                 ...                     ...   
498    XYL          Xylem Inc.             Industrials   
499    YUM         Yum! Brands  Consumer Discretionary   
500   ZBRA  Zebra Technologies  Information Technology   
501    ZBH       Zimmer Biomet             Health Care   
502    ZTS              Zoetis             Health Care   

                                GICS Sub-Industry    Headquarters Location  \
0                        Industrial Conglomerates    Saint Paul, Minnesota   
1                               Building Products     Milwaukee, Wisconsin   
2                           Health Care Equipment  North Chicago, Illinoi
```

```python
# we can look at the first n rows of a table using df.head(n)
# please do so here
df_head5 = df.head(5)
```

```python
grader.check("q1a")
```

Output:
```
q1a results: All test cases passed!
```

In this analysis, we are not particularly interested in the columns for security, headquarters, date added, founded, or CIK values. We can select specific columns of a DataFrame by passing a list of column names:

```python
columns = ['colname1', 'colname2', ...]
df[columns]
```

Please select every column except for security, headquarters, CIK, founded, and date added.

```python
columns = ["Symbol","GICS Sector", "GICS Sub-Industry"]
df[columns]
```

Output:
```
Symbol             GICS Sector  \
0      MMM             Industrials   
1      AOS             Industrials   
2      ABT             Health Care   
3     ABBV             Health Care   
4      ACN  Information Technology   
..     ...                     ...   
498    XYL             Industrials   
499    YUM  Consumer Discretionary   
500   ZBRA  Information Technology   
501    ZBH             Health Care   
502    ZTS             Health Care   

                                GICS Sub-Industry  
0                        Industrial Conglomerates  
1                               Building Products  
2                           Health Care Equipment  
3                                   Biotechnology  
4                  IT Consulting & Other Services  
..                                            ...  
498  Industrial Machinery & Supplies & Components  
499                                   Restaurants  
500            Electronic Equipment & Instruments  
501                        
```

```python
grader.check("q1b")
```

Output:
```
q1b results: All test cases passed!
```

It appears there are quite a few sectors! In a DataFrame, one column is referred to as a series. We can select a series using `df['colname']`. By using `series.unique()`, we can obtain a list of the unique values within a series.

Please return the number of unique values in the 'GICS Sector' column. Note that you can programmatically access the length of this array using `.size`.

```python
unique_sectors = len(df["GICS Sector"].unique())
unique_sectors
```

Output:
```
11
```

```python
grader.check("q1c")
```

Output:
```
q1c results: All test cases passed!
```

If we want to determine how many times each unique value appears in a series, we can use `series.value_counts()`. This function associates each unique value with the number of times it appears in the series.

Please return the number of times 'Information Technology' appears in the 'GICS Sector' column. Note that you can programmatically access the count for 'Information Technology' using `series.value_counts()['uniquevalname']`.

```python
it_sectors = int(df["GICS Sector"].value_counts()["Information Technology"])
```

```python
grader.check("q1d")
```

Output:
```
q1d results: All test cases passed!
```

One way to select values from a DataFrame is by using index-based location, or `iloc`. We can use `df.iloc[index for rows, index for columns]` to specify which indices to select.

- `df.iloc[:,:]` will return the entire DataFrame.
- `df.iloc[:5, :6]` will return the first 5 rows and 6 columns.
- `df.iloc[2:5, 3:]` will return the third column and onward, for rows 2-4.

Please return rows 5-9, and return the 2nd column and onward.

```python
specified_slice = df.iloc[5:10, 2:]
specified_slice
```

Output:
```
GICS Sector                             GICS Sub-Industry  \
5  Information Technology                          Application Software   
6  Information Technology                                Semiconductors   
7               Utilities  Independent Power Producers & Energy Traders   
8              Financials                       Life & Health Insurance   
9             Health Care                Life Sciences Tools & Services   

     Headquarters Location  Date added      CIK Founded  
5     San Jose, California  1997-05-05   796343    1982  
6  Santa Clara, California  2017-03-20     2488    1969  
7      Arlington, Virginia  1998-10-02   874761    1981  
8        Columbus, Georgia  1999-05-28     4977    1955  
9  Santa Clara, California  2000-06-05  1090872    1999
```

```python
grader.check("q1e")
```

Output:
```
q1e results: All test cases passed!
```

We can also examine the sub-industries, each of which is associated with one of the sectors. To select only the rows of the DataFrame that correspond to a specific sector, we can use boolean indexing.

A boolean operator can be created generally using `df['colname'] {comparator} {value}`. This returns a series of `False` and `True` values for each row.

Please create a boolean operator for when the 'GICS Sector' is 'Health Care' and return the number of `False` values using `value_counts()`.

```python
false_healthcare = int((df["GICS Sector"] == "Health Care").value_counts()[False])
false_healthcare
```

Output:
```
439
```

```python
grader.check("q1f")
```

Output:
```
q1f results: All test cases passed!
```

Now that we have our boolean operator, we can select the values that are `True`. We do this using `df[operator]`, where `operator = df['colname'] {comparator} {value}`.

Please use the operator you created to select only the rows where the security is in the 'Energy' sector (using the 'GICS Sector' column).

```python
energy = df[df["GICS Sector"] == "Energy"]
```

```python
grader.check("q1g")
```

Output:
```
q1g results: All test cases passed!
```

We also have `df.loc`, which works similarly to `df.iloc`. Here are some examples:

- `df.loc[row_indexer, column_indexer]`
- `df.loc[start_row_label:end_row_label, start_col_label:end_col_label]`
- `df.loc[list_of_row_labels, list_of_column_labels]`
- `df.loc[boolean_array_for_rows, boolean_array_for_columns]`

Please select the rows where the 'GICS Sector' is 'Industrials', and return the 1st through 4th columns.

```python
first4cols_industrials = df[df["GICS Sector"] == "Industrials"].iloc[:, :4]
```

```python
grader.check("q1h")
```

Output:
```
q1h results: All test cases passed!
```

Another useful DataFrame operation is sorting. We can sort an entire DataFrame using `df.sort_values('colname', ascending=True/False)`.

Please sort, in ascending order, the symbols in the 'Consumer Discretionary' sector, and report the 4th oldest security (by date added). You may want to break this into three separate tasks to make it easier.

```python
old_4_security = (df[df["GICS Sector"] == "Consumer Discretionary"].sort_values("Date added", ascending=True)).iloc[3,1]
old_4_security
```

Output:
```
'Bath & Body Works, Inc.'
```

```python
grader.check("q1i")
```

Output:
```
q1i results: All test cases passed!
```

Let's also introduce the `groupby` operator. We can group a DataFrame based on a given column, such as the 'GICS Sector' column. After grouping a DataFrame, we commonly apply some kind of aggregation, such as a count, an average, etc. For example, we might use `df.groupby('colname').count()`.

Please return the average CIK number of the 'Real Estate' sector using the `df.groupby('colname').mean(numeric_only=True)` aggregation. You can also access a specific series by adding `df.groupby('colname1').mean(numeric_only=True)['colname2']` to view the aggregation of a particular series.

```python
mean_realestate_cik = df[df["GICS Sector"] == "Real Estate"].groupby("GICS Sector").mean(numeric_only=True)["CIK"].reset_index().iloc[0,1]
mean_realestate_cik
```

Output:
```
np.float64(956436.7741935484)
```

```python
grader.check("q1j")
```

Output:
```
q1j results: All test cases passed!
```

```python
def get_common_subindustry(df):
    # .index returns the index of a series, in this case the unique values
    # since value counts is sorted in descending order
    # .value_counts().index[0] will give us the largest avlue
    
    # we can also use display to show the intermediate steps
    # we use series.values to access the list of values in the series
    
    print('-'*35)
    print(df['GICS Sector'].values[0])
    print('-'*35)
    display(df['GICS Sub-Industry'].value_counts())
    return df['GICS Sub-Industry'].value_counts().index[0]
```

```python
df['GICS Sub-Industry'].value_counts().index[0]
```

Output:
```
'Health Care Equipment'
```

```python
df.groupby('GICS Sector').apply(get_common_subindustry)
```

Output:
```
-----------------------------------
Communication Services
-----------------------------------
GICS Sub-Industry
Interactive Media & Services             4
Movies & Entertainment                   4
Broadcasting                             3
Cable & Satellite                        2
Integrated Telecommunication Services    2
Interactive Home Entertainment           2
Advertising                              2
Publishing                               2
Wireless Telecommunication Services      1
Name: count, dtype: int64
-----------------------------------
Consumer Discretionary
-----------------------------------
GICS Sub-Industry
Hotels, Resorts & Cruise Lines         8
Restaurants                            6
Homebuilding                           4
Casinos & Gaming                       4
Apparel, Accessories & Luxury Goods    4
Broadline Retail                       3
Other Specialty Retail                 3
Distributors                           3
Automobile Manufacturers         
```

# data downloading/cleaning

```python
# the yfinance library provides access to lots of information about tickers available on yahoo finance
!pip install yfinance
```

Output:
```
Collecting yfinance
  Using cached yfinance-0.2.55-py2.py3-none-any.whl.metadata (5.8 kB)
Requirement already satisfied: pandas>=1.3.0 in /srv/conda/lib/python3.11/site-packages (from yfinance) (2.2.3)
Requirement already satisfied: numpy>=1.16.5 in /srv/conda/lib/python3.11/site-packages (from yfinance) (2.2.4)
Requirement already satisfied: requests>=2.31 in /srv/conda/lib/python3.11/site-packages (from yfinance) (2.32.3)
Collecting multitasking>=0.0.7 (from yfinance)
  Using cached multitasking-0.0.11-py3-none-any.whl.metadata (5.5 kB)
Requirement already satisfied: platformdirs>=2.0.0 in /srv/conda/lib/python3.11/site-packages (from yfinance) (4.3.6)
Requirement already satisfied: pytz>=2022.5 in /srv/conda/lib/python3.11/site-packages (from yfinance) (2024.1)
Requirement already satisfied: frozendict>=2.3.4 in /srv/conda/lib/python3.11/site-packages (from yfinance) (2.4.6)
Collecting peewee>=3.16.2 (from yfinance)
  Using cached peewee-3.17.9-cp311-cp311-linux_x86_64.whl
Requireme
```

```python
import yfinance as yf
```

**Errors are expected as data is missing for some tickers, see the future cells!**

```python
# we can download historical price data for these assets using
# yf.download(tickers)
# there's other optional parameters that we won't cover here
# lets get the list of symbols/tickers we care about, and then download them using .unique() on our dataframe
# convert the series to a list before you pass it to yf
# this may take a minute or two
# please put the data into a dataframe called hdf, or historical dataframe for short

tickers = list(df['Symbol'].unique())
hdf = yf.download(tickers, start='2023-10-01', end='2024-03-29')
hdf = hdf[hdf.index <= '2024-03-19']
```

Output:
```
[*********************100%***********************]  503 of 503 completed

8 Failed downloads:
['BRK.B', 'WRK', 'PXD', 'CTLT', 'MRO']: YFTzMissingError('possibly delisted; no timezone found')
['SOLV', 'GEV']: YFPricesMissingError('possibly delisted; no price data found  (1d 2023-10-01 -> 2024-03-29) (Yahoo error = "Data doesn\'t exist for startDate = 1696132800, endDate = 1711684800")')
['BF.B']: YFPricesMissingError('possibly delisted; no price data found  (1d 2023-10-01 -> 2024-03-29)')
```

```python
grader.check("q2a")
```

Output:
```
q2a results: All test cases passed!
```

Let's take a look at the data we've just downloaded

```python
hdf
```

Output:
```
Price      Adj Close                                       Close         ...  \
Ticker          BF.B BRK.B CTLT GEV MRO PXD SOLV WRK           A    AAL  ...   
Date                                                                     ...   
2023-10-02       NaN   NaN  NaN NaN NaN NaN  NaN NaN  109.716743  12.75  ...   
2023-10-03       NaN   NaN  NaN NaN NaN NaN  NaN NaN  109.083572  12.29  ...   
2023-10-04       NaN   NaN  NaN NaN NaN NaN  NaN NaN  110.478531  12.73  ...   
2023-10-05       NaN   NaN  NaN NaN NaN NaN  NaN NaN  109.172607  12.85  ...   
2023-10-06       NaN   NaN  NaN NaN NaN NaN  NaN NaN  109.459526  12.76  ...   
...              ...   ...  ...  ..  ..  ..  ...  ..         ...    ...  ...   
2024-03-13       NaN   NaN  NaN NaN NaN NaN  NaN NaN  146.490189  14.00  ...   
2024-03-14       NaN   NaN  NaN NaN NaN NaN  NaN NaN  143.903656  13.99  ...   
2024-03-15       NaN   NaN  NaN NaN NaN NaN  NaN NaN  146.153229  13.96  ...   
2024-03-18       NaN   NaN  NaN NaN NaN 
```

```python
# looks like the index is the date, and we have multiple levels of columns!
# we can get the top level of columns using df.columns.get_level_values(i)
# let's look at the unique top level columns, with .unique()
```

```python
hdf.columns.get_level_values(0).unique()
```

Output:
```
Index(['Adj Close', 'Close', 'High', 'Low', 'Open', 'Volume'], dtype='object', name='Price')
```

```python
# the first level of columns has the 'adjusted close', or the closing price for that day, with interest rate adjusted
# we are also given 'volume', or the amount of shares traded that day
# the other stuff might be nice to have, but we'll just stick with adjusted close and volume for simplicity
# let's now select those
# to not destroy our old dataframe, we'll make a copy called hdf_raw with df.copy()
# once you make the copy, make hdf = hdf[columns]
hdf_raw = hdf.copy()
hdf = hdf[['Adj Close', 'Volume']]
```

2b or not 2b, the answer is not because I deleted the question :) Also, **the next cell is filled in for you, so there's nothing you need to do here.**

```python
# you may have seen the following error message upon download: 

# 2 Failed downloads:
# - BF.B: No data found for this date range, symbol may be delisted
# - BRK.B: No data found, symbol may be delisted
# - GEV and SOLV are also recent additions without any data

# please use df.loc to view the data for these symbols
# you should return all rows, and use columns.get_level_values(1) to get the second level of columns, with the symbols
# note that you can use '.isin(list)' as a boolean operator to return if a value is in a list
# your code might look like

# df.loc[:,df.columns.get_level_values(1).isin([list of values to check])]
broken_df = hdf.loc[:,hdf.columns.get_level_values(1).isin(['BRK.B', 'BF.B', 'GEV', 'SOLV', 'PXD', 'WRK', 'CTLT', 'MRO'])]
broken_df
```

Output:
```
Price      Adj Close                                 Volume                 \
Ticker          BF.B BRK.B CTLT GEV MRO PXD SOLV WRK   BF.B BRK.B CTLT GEV   
Date                                                                         
2023-10-02       NaN   NaN  NaN NaN NaN NaN  NaN NaN    NaN   NaN  NaN NaN   
2023-10-03       NaN   NaN  NaN NaN NaN NaN  NaN NaN    NaN   NaN  NaN NaN   
2023-10-04       NaN   NaN  NaN NaN NaN NaN  NaN NaN    NaN   NaN  NaN NaN   
2023-10-05       NaN   NaN  NaN NaN NaN NaN  NaN NaN    NaN   NaN  NaN NaN   
2023-10-06       NaN   NaN  NaN NaN NaN NaN  NaN NaN    NaN   NaN  NaN NaN   
...              ...   ...  ...  ..  ..  ..  ...  ..    ...   ...  ...  ..   
2024-03-13       NaN   NaN  NaN NaN NaN NaN  NaN NaN    NaN   NaN  NaN NaN   
2024-03-14       NaN   NaN  NaN NaN NaN NaN  NaN NaN    NaN   NaN  NaN NaN   
2024-03-15       NaN   NaN  NaN NaN NaN NaN  NaN NaN    NaN   NaN  NaN NaN   
2024-03-18       NaN   NaN  NaN NaN NaN NaN  NaN NaN    NaN   Na
```

```python
grader.check("q2c")
```

Output:
```
q2c results: All test cases passed!
```

```python
# awesome, now lets select everything except for that, since we dont want to use this data
# we can invert a boolean indexer by using ~, eg. df.loc[:, ~<condition>]
# let hdf be equal to this new selection, and then drop nan values using df.dropna()
hdf = hdf.loc[: , ~hdf.columns.get_level_values(1).isin(['BRK.B', 'BF.B', 'GEV', 'SOLV', 'PXD', 'WRK', 'CTLT', 'MRO'])]
hdf = hdf.dropna()
```

```python
grader.check("q2d")
```

Output:
```
q2d results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Please also check gradescope for any written assignments for this week.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1a (points: 1)
```
>>> assert len(df_head5) == 5
```


Tests for q1b (points: 1)
```
>>> assert len(columns) == 3
```


Tests for q1c (points: 1)
```
>>> assert unique_sectors == 11
```


Tests for q1d (points: 1)
```
>>> assert it_sectors == 65
```
```
>>> it_sectors
65
```


Tests for q1e (points: 1)
```
>>> assert specified_slice.shape == (5,6)
```


Tests for q1f (points: 1)
```
>>> assert false_healthcare == 439
```


Tests for q1g (points: 1)
```
>>> assert energy['CIK'].sum() == 19373487
```


Tests for q1h (points: 1)
```
>>> assert first4cols_industrials.shape == (79,4)
```


Tests for q1i (points: 1)
```
>>> assert old_4_security == 'Bath & Body Works, Inc.'
```


Tests for q1j (points: 1)
```
>>> assert np.isclose(mean_realestate_cik, 956436.7741935484)
```


Tests for q2a (points: 1)
```
>>> assert hdf.shape == hdf.shape
```


Tests for q2b (points: 1)
```
>>> assert hdf.columns.size <= 1006
```


Tests for q2c (points: 1)
```
>>> assert broken_df.shape[1] >= 12
```


Tests for q2d (points: 1)
```
>>> assert hdf.shape[0] == 115 and hdf.shape[1] <= 500
```
